
# Intelligent Audio and Music Analysis Assignment 0

The main purpose of this assignment is to get you familiarized with the server and python infrastructure and workflow.  

## Task 1: Import numpy

Load the python math library [numpy](https://numpy.org/doc/stable/) into namespace variable `np`.

<!-- grade_id: import_numpy, points: 0, grade: false, solution: true -->
```python
# load the python math library numpy into a local namespace variable called "np"
# this is achieved using the import <lib> as <name> syntax

### BEGIN SOLUTION
import numpy as np
### END SOLUTION
```

<!-- grade_id: import_numpy_test, points: 1, grade: true, solution: false -->
```python
# test if numpy has been loaded successfully

np.version.version
print('All tests successful!')
```

Output:
```
All tests successful!
```

## Task 2: Create a Matrix

Create a matrix ([ndarray](https://numpy.org/doc/stable/reference/generated/numpy.ndarray.html)) `x` of random numbers 
with size 4x4.

<!-- grade_id: rand_x, points: 0, grade: false, solution: true -->
```python
# numpy has a module called "random", use the rand function. 
# See documentation at: 
# https://numpy.org/doc/stable/reference/random/generated/numpy.random.rand.html

### BEGIN SOLUTION
x = np.random.rand(4, 4)
### END SOLUTION

print("x = {}".format(x))
```

Output:
```
x = [[0.02914997 0.46028399 0.84311286 0.49864209]
 [0.15219914 0.61744741 0.6310644  0.8592101 ]
 [0.06761009 0.20346184 0.67215242 0.30664995]
 [0.86186699 0.57687905 0.42342323 0.59758039]]
```

<!-- grade_id: rand_x_test, points: 1, grade: true, solution: false -->
```python
# test x for properties:

assert x.shape == (4, 4)
assert all([0 <= ex <= 1 for ex in x.flatten()])
### BEGIN HIDDEN TESTS
assert all([x.flatten()[idx1] != x.flatten()[idx2] for idx1 in range(4*4-1) for idx2 in range(idx1+1, 4*4)])
### END HIDDEN TESTS
print('All tests successful!')
```

Output:
```
All tests successful!
```

<!-- grade_id: rand_x_test_exc, points: 1, grade: true, solution: false -->
```python
# test for exception: 

try:
    x[10, 10]
except IndexError:
    pass
else:       
    raise AssertionError('did not raise ValueError')
print('All tests successful!')
```

Output:
```
All tests successful!
```

## Task 3: Elementwise Matrix Operations
Calculate the *elementwise* square of `x` into a new variable named `y`.

<!-- grade_id: x_mult, points: 0, grade: false, solution: true -->
```python
# Arrays can be used as matrices, thus we can choose between matrix multiplication and elementwise multiplication.
# https://numpy.org/doc/stable/reference/generated/numpy.matmul.html
# https://numpy.org/doc/stable/reference/generated/numpy.multiply.html

### BEGIN SOLUTION
y = np.multiply(x, x)
### END SOLUTION
```

<!-- grade_id: x_mult_test, points: 1, grade: true, solution: false -->
```python
# Test shape of y
# Note: for full points, values of y must also be correct.

assert y.shape == (4, 4)

### BEGIN HIDDEN TESTS
assert np.all(y == np.multiply(x, x))
### END HIDDEN TESTS
print('All tests successful!')
```

Output:
```
All tests successful!
```

## Task 4: Freeform Answers

Sometimes you will find questions that require a textual answer. Use the markdown input field below to submit your answer.
 
Answer the following question:

**What is your background or experience with (Python) programming, audio processing, and machine learning?**